Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

---

# Using classes in Python

## A. Create class for 3D vector 

Create a class for storing a 3D vector (x, y, z coordinates), and which provides a set of functions for manipulating the coordinates.  

* Implement class Vector3D which has the following properties:
* Attributes
    - coord_x stores x coordinate (object attribute)
    - coord_y stores y coordinate (object attribute)
    - coord_z stores z coordinate (object attribute)
* Ordinary class methods 
    - length() : computes length of vector* 
    - norm()   : renormalizes vector to unit length vector*
      * for square root use math.sqrt function from math library
* Overloaded operations  
    - addition operator(+)     : implement `__add__` method 
    - substraction operator(-) : implement `__sub__` method 

<!-- grade_id: cell-106104681ce80401, points: 0, grade: false, solution: true -->
```python
import math

class Vector3D():
    """
    Class for manipulation of 3D-vectors
    
    instance attributes:
        coord_x, coord_y, coord_z (float)
    
    class methods
        length (float): returns Euclidean norm of vector
        normalize (None): scales coordinates so that resulting vector has length 1
        __add__ (Vector3D) returns new vector as vector sum of two vectors
        __sub__ (Vector3D) returns new vector as vector difference of two vectors
        __repr__: string representation of vector

    """
    def __init__(self, x, y, z):
        self.coord_x = x
        self.coord_y = y
        self.coord_z = z

    def length(self):
        self.len = math.sqrt(self.coord_x**2 + self.coord_y**2 + self.coord_z**2)
        return(self.len)

    def normalize(self):
        base = self.length()
        self.coord_x = self.coord_x/base
        self.coord_y = self.coord_y/base
        self.coord_z = self.coord_z/base

    def __add__(self, other):
        return Vector3D(self.coord_x + other.coord_x, self.coord_y + other.coord_y, self.coord_z + other.coord_z)

    def __sub__(self, other):
        return Vector3D(self.coord_x - other.coord_x, self.coord_y - other.coord_y, self.coord_z - other.coord_z)

    def __repr__(self):
        return(f"[{self.coord_x} , {self.coord_y} , {self.coord_z}]")
```

```python
# Define helper function
def assert_close(x, y, max_difference=1e-8):
    calculated_difference = abs(x - y)
    error_message = f'Not close enough: {calculated_difference} > {max_difference}'
    assert calculated_difference < max_difference, error_message
```

<!-- grade_id: cell-e8becee65f92a575, points: 1, grade: true, solution: false -->
```python
# Verify Vector3D constructor
vector = Vector3D(1.0, 2.0, 4.0)
assert_close(vector.coord_x, 1.0)
assert_close(vector.coord_y, 2.0)
assert_close(vector.coord_z, 4.0)
```

<!-- grade_id: cell-3328f70c25343894, points: 1, grade: true, solution: false -->
```python
# Verify length method in Vector3D class 
vector = Vector3D(1.0, 2.0, 4.0)

assert_close(vector.length(), math.sqrt(21.0))
```

<!-- grade_id: cell-96e5a932f3347dcf, points: 1, grade: true, solution: false -->
```python
# Verify norm method in Vector3D class 
vector = Vector3D(1.0, 2.0, 4.0)

vector.normalize()

assert_close(vector.length(),1.0)
```

<!-- grade_id: cell-1f019ebe8894fb2b, points: 1, grade: true, solution: false -->
```python
#Verify addition operator for two vectors 
avector = Vector3D(1.0, 2.0, 4.0)
bvector = Vector3D(2.0, 3.3, 0.9)

cvector = avector + bvector

assert_close(cvector.coord_x, 3.0)
assert_close(cvector.coord_y, 5.3)
assert_close(cvector.coord_z, 4.9)
```

<!-- grade_id: cell-965fcb9cfdf9230b, points: 1, grade: true, solution: false -->
```python
#Verify substraction operator for two vectors 
avector = Vector3D(1.0, 2.0, 4.0)
bvector = Vector3D(2.0, 0.2, 0.9)

cvector = avector - bvector

assert_close(cvector.coord_x, -1.0)
assert_close(cvector.coord_y,  1.8)
assert_close(cvector.coord_z,  3.1)

# MY OWN TEST
print(repr(avector))
```

Output:
```
[1.0 , 2.0 , 4.0]
```

## B. Create class for currency data storage 

Create a class for storing money in SEK or Euros, which provides a set of functions for manipulating the amount and type of currency. 

* Implement class Money which has the following properties:
* Attributes
  - conversion_rate stores conversion rate from EURO to SEK (class attribute, default value 10.0)
  - currency stores currency type ('SEK' or 'EURO') (instance attribute)
  - amount stores amount of money (instance attribute)
* Ordinary class methods
  - set_conversion_rate(new_conversion_rate) : for setting conversion rate 
* Overloaded operations
  - addition operator(+)     : implement addition of money ammounts* 
  - substraction operator(-) : implement substraction of money ammounts* 
  * currency type of result must be currency type of first operand in arithmetical expresion 

<!-- grade_id: cell-dbe1d3027a05e5a4, points: 0, grade: false, solution: true -->
```python
class Money:
    """
    A class for manipulation of different currencies
    
    class attributes:
        conversion_rate (float) 
    instance attributes:
        amount (float)
        currency (str)
    class methods:
        set_conversion_rate: resets conversion rate
        __add__: implements addition of monies (right side currency converted to left side if different)
        __sub__: implements subtraction of monies
    """

    # CONVERSION RATE (CLASS ATTRIBUTE)
    conversion_rate = 10

    def __init__(self, m, c):
        self.amount = m
        self.currency = c
    def set_conversion_rate(self, rate):
        Money.conversion_rate = rate
    def reset_conversion_rate(self):
        Money.conversion_rate = 10
    def __add__(self, other):
        if self.currency == other.currency:
            curr = self.currency
            return Money(self.amount + other.amount, curr)
        else:
            if self.currency == 'SEK':
                curr = self.currency
                other.currency = 'SEK'
                other.amount = other.amount * 10
                return Money(self.amount + other.amount, curr)
            else:
                curr = self.currency
                other.currency = 'EUR'
                other.amount = other.amount // 10
                return Money(self.amount + other.amount, curr)
                # I could've moved the conversion "* 10" or "// 10" into the return and NOT change the specifications of the variable.
                # Wity my method of choice I am changing the specifications of the variables as before returning a new variable.
                # Another method is used below for "__sub__"
    def __sub__(self, other):
        if self.currency == other.currency:
            curr = self.currency
            return Money(self.amount - other.amount, curr)
        else:
            if self.currency == 'SEK':
                curr = self.currency
                return Money(self.amount - (other.amount * 10), curr)
            else:
                curr = self.currency
                return Money(self.amount - (other.amount // 10), curr)
```

<!-- grade_id: cell-73e7fae89242df32, points: 1, grade: true, solution: false -->
```python
#Verify class attribute
assert Money.conversion_rate == 10.0
```

<!-- grade_id: cell-4181fe5f998c9a96, points: 1, grade: true, solution: false -->
```python
#Verify instance attributes
ma = Money(100, 'SEK')
mb = Money(50, 'EURO')
assert ma.amount == 100
assert ma.currency == 'SEK'
assert mb.amount == 50
assert mb.currency == 'EURO'
```

<!-- grade_id: cell-fa09de3f68f4fad3, points: 1, grade: true, solution: false -->
```python
#Verify addition
ma = Money(100, 'SEK')
mb = Money(50, 'EURO')
mc = ma + mb
assert mc.amount == 600
assert mc.currency == 'SEK'
```

<!-- grade_id: cell-77d09f62ca89cb1d, points: 1, grade: true, solution: false -->
```python
#Verify subtraction
ma = Money(100, 'SEK')
mb = Money(50, 'EURO')
md = mb - ma
assert md.amount == 40
assert md.currency == 'EURO'
```

<!-- grade_id: cell-c2d199586e7c4cbb, points: 1, grade: true, solution: false -->
```python
#Set conversion rate
ma = Money(100, 'SEK')
mb = Money(50, 'EURO')
ma.set_conversion_rate(20)
assert ma.conversion_rate == 20
assert mb.conversion_rate == 20
```

## C. Animals

Complete the Animal class to have and a class variable `animals` (list) and two instance variables, `name` (str) and `number` (int). You need to implement `__init__` and `__str__` methods

<!-- grade_id: cell-f9faff46e9639b35, points: 0, grade: false, solution: true -->
```python
class Animal:
    """
    A class for storing animals
    
    class attributes:
        animals: (list) to store all animals
    instance attributes:
        name:  (str) to store animal name
        number: (int) to store animal order number (starting with 1)
        
    class methods:
        __str__: string representation of animal, e.g. "1. Dog"
        
    static methods:
        zoo: returns string representation of all animals in orderd lies, e.g. 
           '''
           1. Dog
           2. Cat'''
    """
    # ANIMALS (CLASS ATTRIBUTE)
    animals = []

    def __init__(self, name):
        # Appending names to animals
        self.name = name

        # Animal.animals_name.append(name)
        Animal.animals.append(self)

        # Assigning numbers to animals
        self.number = len(Animal.animals)

    def __str__(self):
        namecap = self.name.title()
        return(f'{self.number}. {namecap}')
        
    @staticmethod
    def zoo():
        zoo_list = []
        for i,j in zip(Animal.animals, Animal.animals):
            zoo_list.append(f'{i.number}. {j.name.title()}')
        return('\n'.join(zoo_list))
```

<!-- grade_id: cell-41580406aa20538d, points: 1, grade: true, solution: false -->
```python
Animal.animals.clear()

dog = Animal('dog')
assert dog.name == 'dog'
assert dog.number == 1
assert str(dog) == '1. Dog'

cat = Animal('cat')
assert cat.name == 'cat'
assert cat.number == 2
assert str(cat) == '2. Cat'

```

A static method is a function in a class that is like an ordinary function, that does not depend on any instance (no self argument). 

In a class it is defined with a `@staticmethod` decorator.

It can be appended to the class definition as below. Complete the static method so that it returns a string which lists all member animals

<!-- grade_id: cell-a89d339b99865f5d, points: 1, grade: true, solution: false -->
```python
Animal.animals.clear()

#Generate a list of animals and compare with the class attribute
animals = [Animal(a) for a in ['camel', 'donkey', 'hippo']]
assert animals == Animal.animals, f'{animals} != {Animal.animals}'

#zoo should produce a printout of the defined animals
zoo_output = Animal.zoo()
print(zoo_output)
expected_output = """
1. Camel
2. Donkey
3. Hippo
"""

condition = zoo_output.strip() == expected_output.strip()
error_message = f"\n{zoo_output}\n   !=\n{expected_output}"

assert condition, error_message
```

Output:
```
1. Camel
2. Donkey
3. Hippo
```

# PAST EXAM QUESTIONS:

### Car - `basics`, `repr`, `str`, `list creation`, `sorting`, `inherit`, etc. 

#### Car - Basics (`self.XX`)

Outline a class definition Car for a car with attributes make, model, year.

Solution:

```python
class Car:
    def __init__(self, make, model, year):
        self.make = make
        self.model = model
        self.year = year
```

#### Car continued - String representation (`__repr__`)
Add a class method so that default string representation of an object mimics the command for creation.

Such that:

    >>> car = Car('Volvo', 'Amazon', 1964)
    >>> repr(car)
    "Car('Volvo', 'Amazon', 1964)"

Solution:

```python
class Car:
    def __init__(self, make, model, year):
        self.make = make
        self.model = model
        self.year = year
        
    def __repr__(self):
        return f"Car('{self.make}', '{self.model}', {self.year})"
```

#### Car continued - Another string read (`__str__`)
Add another class method so that the other string representation of an object reads

Such that:

    >>> car = Car('Volvo', 'Amazon', 1964)
    >>> str(car)
    'Volvo Amazon (1964)'

Solution:

```python
class Car:
    def __init__(self, make, model, year):
        self.make = make
        self.model = model
        self.year = year
    
    def __repr__(self):
        return f"{self.__class__.__name__}('{self.make}', '{self.model}', {self.year})"
        
    def __str__(self):
        return f"{self.make} {self.model} ({self.year})"
```

#### Car continued - Generating list by reading a csv file

Outline code that generate a list of Car objects by reading the a file `cars.csv` containing ...

Dodge,Charger,1969 <br>
GMC,Vandura G2500,1995 <br>
Toyota,Sienna,2007 <br>
Dodge,Challenger,2012 <br>
Pontiac,Grand Am,1989 <br>
Nissan,Altima,2009 <br>
Mazda,MPV,2002 <br>
Cadillac,DeVille,1994 <br>
Mercury,Tracer,1999 <br>
Volkswagen,Passat,1988 <br>

Such that:

    >>> cars
    [Car('Dodge', 'Charger', 1969), Car('GMC', 'Vandura G2500', 1995), Car('Toyota', 'Sienna', 2007), Car('Dodge', 'Challenger', 2012), Car(...) ...]

Solution:

```python
def carlist(filename):
    for line in open(filename):
        make, model, year = line.strip().split(',')
        car = Car(make, model, year)
        cars.append(car)

filename = 'cars.csv'
```

#### Car continued - Sorting

Sort the cars by year from newest to oldest.

The `sorted` function has the documentation

`sorted(iterable, /, *, key=None, reverse=False)
Return a new list containing all items from the iterable in ascending order.`

`A custom key function can be supplied to customize the sort order, and the
reverse flag can be set to request the result in descending order.`

How can you use this to sort the cars by year from newest to oldest such that:

    >>> sorted(cars, key=get_year, reverse=True)
    [Car('Dodge', 'Challenger', 2012), Car('Nissan', 'Altima', 2009), Car('Toyota', 'Sienna', 2007), Car('Mazda', 'MPV', 2002), Car('...), ...]

```python
def get_year(car):
    return car.year
```

#### Car continued - Inherit class

Define a new class that inherits from Car with an initial price zero

A car salesman wants to use your code but update to have a price attribute. Define a new class that inherits from Car with an initial price zero, such that:

    >>> car = Car4Sale('Volvo', 'Amazon', 1964)
    >>> car.price
    0

    >>> car = Car4Sale('Volvo', 'Amazon', 1964, 9900)
    >>> car.price
    9900

Solution:

```python
class Car4Sale(Car):
    def __init__(self, make, model, year, price=0):
        super().__init__(make, model, year)
        self.price = price
```

#### Car continued - Calculate total price
Write the function to calculate the total price for a list of cars

Such that:

    >>> sum_values([])
    0
    >>> sum_values([
    ... Car4Sale('Chevrolet' ,'Silverado 3500',2003,34452),
    ... Car4Sale('Mazda', 626, 1991, 17121),
    ... Car4Sale('Oldsmobile', 'Achieva', 1993, 12982),
    ... ])
    64555

Solution:

```python
def sum_values(cars):
    return sum(car.price for car in cars)
```

#### Car continued - Composition

When it comes to extending a class an alternative to inheritance is so called composition...

When it comes to extending a class an alternative to inheritance is so called composition, which means in this case that a car and its price are
separate data attributes of a new class:

    >>> class CarWithPrice:
    ... def __init__(self, car, price=0):
    ... self.car = car
    ... self.price = price
    ...
    ... def __str__(self):
    ... return f"{self.car}: {self.price}"
    >>> car = Car('Mercury', 'Sable', 1988)
    >>> car_price = CarWithPrice(car, 7000)
    
What will be the output of `print(car_price)`?

Solution:

    >>> print(car_price)
    Mercury Sable (1988): 7000

### What is the name of the ``class method`` where instance variables are defined?

In a class defintion it is:

`__init__`

### Calling a ``class`` in multiple ways

A class Foo containing the following definition of hi could potentially be
called the following way.

    >>> class Foo:
    ... def hi(self):
    ... print(’hi’)
    >>> foo = Foo()
    >>> Foo.hi(foo)
    hi

What does a normal class method call look like in this case?

Answer:

    >>> foo.hi()
    hi