Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

---

# File handling


## Creating our reference file

In our initial example we create a small reference file by executing the following cell. The so-called cell-magic command `%%file` takes its first argument as a filename and writes the remaining content of the cell to file

[More about cell magic](https://ipython.readthedocs.io/en/stable/interactive/magics.html)

```python
%%file composers.txt
dufay
ockeghem
josquin
willaert
orlando
sweelinck
```

## Reading

In the first exercise we recall different forms of reading text files. These are single-line answers that output some text. 



#### Exercise 1
What single line is used to return the file content as a single string? Not that the string will contain the end-of-line (newline) character `\n`

<!-- grade_id: cell-619ee4026f28565c, points: 0, grade: false, solution: true -->
```python
def text_to_string():
    data = open("composers.txt").read() # This single line is used to reutn the file content as a single string.
    # print(data)
    return(data)
text_to_string()
```

An underscore `_` refers to the value of the previoius command. In the test cell we are thus checking the return  of the last line in the previous cell

<!-- grade_id: cell-e16f081e726162ce, points: 1, grade: true, solution: false -->
```python
result = _
expected = 'dufay\nockeghem\njosquin\nwillaert\norlando\nsweelinck\n'
assert result == expected, f'{result=} != {expected}'
```

#### Exercise 2
What single line is used to return the file content as a Python list of strings. 

<!-- grade_id: cell-b8288e222d0343fb, points: 0, grade: false, solution: true -->
```python
def text_to_list():
    data = open("composers.txt").readlines() # This single line is used to reutn the file content as a list.
    # print(data)
    return(data)
text_to_list()
```

<!-- grade_id: cell-3f4d59f2de6a1c5e, points: 1, grade: true, solution: false -->
```python
result = _
expected = ['dufay\n', 'ockeghem\n', 'josquin\n',  'willaert\n',  'orlando\n',  'sweelinck\n']
assert result == expected, f'{result=} != {expected}'
```

#### Exercise 3
This function will return the lines of the file without the 'end-of-line' characters. Use the [`strip`](https://docs.python.org/3/library/stdtypes.html#str.strip) method

<!-- grade_id: cell-a53c2938a84d15db, points: 0, grade: false, solution: true -->
```python
def get_list_of_lines_without_eol_character(filename):
    lines = []
    data = open("composers.txt").readlines()
    for i in data:
        lines.append(i.strip('\n'))
    return lines
 
get_list_of_lines_without_eol_character('composers.txt')
```

<!-- grade_id: cell-de85c7e2aeabc67f, points: 1, grade: true, solution: false -->
```python
result =  get_list_of_lines_without_eol_character('composers.txt')
expected = ['dufay', 'ockeghem', 'josquin', 'willaert', 'orlando', 'sweelinck']
assert result == expected, f'{result=} != {expected}'
```

## Reading numeric data
All reading/writing is done with text objects, i.e. reading a '7' will be text string with a single character rather than a numeric value. 

To get numbers from file intended for computation we have to do a type conversion (typically `int` or `float`)

Consider a file with one number per line. 

    1
    2
    3
    4


### Exercise 4

Write a function that takes a file name as input and returns the sum of the numbers in the file, assuming one number per line

```python
%%file numbers1.txt
1
2
3
4
5
```

<!-- grade_id: cell-e61d832d811f3cc4, points: 0, grade: false, solution: true -->
```python
def sum_numbers_in_file(filename):
    number_sum = 0
    lines = []
    data = open(filename).readlines()
    for i in data:
        lines.append(i.strip('\n'))
    for i in lines:
        number_sum += int(i)
    return number_sum 
 
sum_numbers_in_file('numbers1.txt')
```

<!-- grade_id: cell-cedf19ef9aeb39f9, points: 1, grade: true, solution: false -->
```python
result = sum_numbers_in_file('numbers1.txt')
expected = 15
assert result == expected, f'{result=} != {expected}'
```

### Exercise 5

It was realized that some files have more numbers on several lines such that for a file like

    1 2 3 4 5
    6 7 8 9 10
    
could not be handled by this function but we would get an error like

    E           ValueError: invalid literal for int() with base 10: '1 2 3 4 5\n'
    

In this exercise generalize the function above by copying here. Use the `split` method to split a line into words.
Example

    >>> '1 2 3 4 5\n'.split()
    ['1', '2', '3', '4', '5']
   

```python
%%file numbers2.txt
1 2 3 4 5
6 7 8 9 10
```

```python
# You may uncomment this line and try
# sum_numbers_in_file('numbers2.txt')

# I VERIFIED THAT IT IS NOT WORKING...
```

In this exercise generalize the function from exercise 4. Use the `split` method to split a line into words.

<!-- grade_id: cell-e61d832d811f3cc5, points: 0, grade: false, solution: true -->
```python
def sum_numbers_in_file2(filename):
    number_sum = 0
    lines = []
    data = open(filename).readlines()
    for i in data:
        lines.append(i.strip('\n'))
    lines = ' '.join(lines).split(' ')
    for i in lines:
        number_sum += int(i)
    return number_sum 
   
sum_numbers_in_file2('numbers2.txt')
```

<!-- grade_id: cell-b4d65b9e6d0f1e01, points: 1, grade: true, solution: false -->
```python
result1 = sum_numbers_in_file2('numbers1.txt')
expected1 = 15
assert result1 == expected1, f'{result1=} != {expected1}'

result2 = sum_numbers_in_file2('numbers2.txt')
expected2 = 55
assert result2 == expected2, f'{result2} != {expected2}'
```

## Writing files

When we call the builtin function `open` with a filename as a single argument the default behaviour is to look for an existing file for reading. If it does not exist the system prints  a `FileNotFoundError` and a message like

    open('no_such_file')
    ---------------------------------------------------------------------------
    FileNotFoundError                         Traceback (most recent call last)
    <ipython-input-15-b99a18bdaf59> in <module>()
    ----> 1 open('no_such_file')

    FileNotFoundError: [Errno 2] No such file or directory: 'no_such_file'
    
 
 
To open a new file to be written we have to supply an additional parameter, `w` for writing.

    open('new_file', 'w')
 
The `open` function returns a file object which has methods for reading (`read`, `readlines`) and writing (`write`)

You can often look up the documentation of a function with

```python
# help(open)
```

### Exercise 6

With the content in the multi-line string write down code that writes the following content to a file of the form first_last.txt  and with capitalized names (First, Last) in the text.

Hint: consider string methods

    >>> 'JOHN'.lower()
    john
    >>> 'JOHN'.capitalize()
    John
    
    

<!-- grade_id: cell-9e5cf46774b3a243, points: 0, grade: false, solution: true -->
```python
def try_out_names(first, last):
    content = f"""
    'Tis but thy name that is my enemy;
    Thou art thyself, though not a {last}.
    What's {last}? It is nor hand, nor foot,
    Nor arm, nor face, nor any other part
    Belonging to a man. O, be some other name!
    What's in a name? That which we call a rose
    By any other word would smell as sweet;
    So {first} would, were he not {first} call'd,
    Retain that dear perfection which he owes
    Without that title. {first}, doff thy name,
    And for that name which is no part of thee
    Take all myself.
    """
    return content

def write_to_file(first, last):
    """
    Writes to file first_last.txt balcony scene with names
    Filename will be in lower case
    Names in text will have initial letter capitalized
    """
    name = '_'.join([first,last]).lower()
    name = '.'.join([name,'txt'])
    content = try_out_names(first.capitalize(),last.capitalize())

    file = open(name,'w').write(content)
    return(file)
    
```

<!-- grade_id: cell-afe447054334d72a, points: 1, grade: true, solution: false -->
```python
# Now write to new file and check content
write_to_file('RONALD', 'SMITH')
result = open('ronald_smith.txt').read()
expected= """
    'Tis but thy name that is my enemy;
    Thou art thyself, though not a Smith.
    What's Smith? It is nor hand, nor foot,
    Nor arm, nor face, nor any other part
    Belonging to a man. O, be some other name!
    What's in a name? That which we call a rose
    By any other word would smell as sweet;
    So Ronald would, were he not Ronald call'd,
    Retain that dear perfection which he owes
    Without that title. Ronald, doff thy name,
    And for that name which is no part of thee
    Take all myself.
    """
assert result == expected, f'{result} != {expected}'

write_to_file('ROMEO', 'MONTAGUE')
result = open('romeo_montague.txt').read()
expected = """
    'Tis but thy name that is my enemy;
    Thou art thyself, though not a Montague.
    What's Montague? It is nor hand, nor foot,
    Nor arm, nor face, nor any other part
    Belonging to a man. O, be some other name!
    What's in a name? That which we call a rose
    By any other word would smell as sweet;
    So Romeo would, were he not Romeo call'd,
    Retain that dear perfection which he owes
    Without that title. Romeo, doff thy name,
    And for that name which is no part of thee
    Take all myself.
    """
assert result == expected, f'{result} != {expected}'
```

### Exercise 7

Writing numeric data to a file implies first converting to a string. The built-in function print does that. Look up the documentation on print how you can use it for writing to a file 

```python
# help(print)
```

<!-- grade_id: cell-df4516027c70782a, points: 0, grade: false, solution: true -->
```python
# Print the number pi to a file with pi.txt using print
from math import pi

with open('pi.txt', 'w') as pi_file:
    print(pi, file = pi_file)
    # pi_file.close()
```

<!-- grade_id: cell-37e29477ff926375, points: 1, grade: true, solution: false -->
```python
# Check the first four characters in the file
assert open('pi.txt').read()[:4] == '3.14'
```

### Exercise 8
Implement a file copy command command with Python read an write commands


<!-- grade_id: cell-6360c6f2039531d7, points: 0, grade: false, solution: true -->
```python
def cp(source, target):
    """
    Copies textfile source to target
    """
    with open(source,'r') as old_file, open(target,'w') as new_file:
        for line in old_file:
            new_file.write(line)
    return(new_file)
```

```python
%%file blake.txt
To see a World in a Grain of Sand
And a Heaven in a Wild Flower,
Hold Infinity in the palm of your hand 
And Eternity in an hour.
```

<!-- grade_id: cell-d7bfacb6d05ced8a, points: 1, grade: true, solution: false -->
```python
cp('blake.txt', 'poem.txt')
result = open('poem.txt').read()
expected = open('blake.txt').read()
assert result == expected, f'{result} != {expected}'
```