## Exercise 09.1 (checking data validity)

The Fibonacci series is valid only for $n \ge 0$. Add to the Fibonacci function in this notebook a check that raises an exception if $n < 0$. Try some invalid data cases to check that an exception is raised.

*Optional:* Use `pytest` to test that an exception *is* raised for some $n < 0$ cases.

<!-- grade_id: cell-982c3611466532a2, points: 0, grade: false, solution: true -->
```python
def f(n):
    if n < 0:
        raise ValueError
        
    if n <= 1:
       return n
    else:
       return(f(n-1) + f(n-2))
        
# n = 'bla'
# f(-1)
```

<!-- grade_id: cell-9b97ea3e85d77f4a, points: 0, grade: true, solution: false -->
```python
# Perform some tests    
assert f(0) == 0
assert f(1) == 1
assert f(2) == 1
assert f(3) == 2
assert f(10) == 55
assert f(15) == 610

# Check that ValueError is raised for n < 0
import pytest
with pytest.raises(ValueError):
    f(-1)
with pytest.raises(ValueError):
    f(-2)
```

## Exercise 09.2 (raising exceptions)

Modify your program from the bisection exercise in Activity 04 to raise an error if the maximum number of iterations is exceeded. Reduce the maximum allowed iterations to test that an exception is raised.

Add any other checks on the input data that you think are appropriate.

<!-- grade_id: cell-ca2c3ee62ad3b956, points: 0, grade: false, solution: true -->
```python
def my_f(x):
    "Evaluate polynomial function"
    return x**3 - 6*x**2 + 4*x + 12

def compute_root(f, x0, x1, tol, max_it):
    
    "Compute roots of a function using bisection"
    error = tol + 1.0

    # Iterate until tolerance is met
    it = 0
    while error > tol:
        
        # Iteration counter
        it += 1
        
        # Guard against an infinite loop
        if it > max_it:
            print("Oops, iteration count is very large. Breaking out of while loop.")
            raise RuntimeError
        
        # Compute midpoint
        x_mid = (x0 + x1)/2

        # Evaluate function at left end-point and at midpoint
        f0 = x0**3 - 6*x0**2 + 4*x0 + 12
        f = x_mid**3 - 6*x_mid**2 + 4*x_mid + 12
        f_mid = f
        
        if f0 * f < 0:
            x1 = x_mid
        else:
            x0 = x_mid
        
        # Computating error
        error = abs(f)
        
        print(it, x_mid, error)
     
    return x_mid, f_mid, it
```

<!-- grade_id: cell-f6cf91ed5dd1736b, points: 0, grade: true, solution: false -->
```python
# Test with max_it = 30
x, f, num_it = compute_root(my_f, x0=3, x1=6, tol=1.0e-6, max_it=30)

# Test with max_it = 20
with pytest.raises(RuntimeError):
    x, f, num_it = compute_root(my_f, x0=3, x1=6, tol=1.0e-6, max_it=20)
```

Output:
```
1 4.5 0.375
2 5.25 12.328125
3 4.875 4.763671875
4 4.6875 1.910888671875
5 4.59375 0.699554443359375
6 4.546875 0.14548873901367188
7 4.5234375 0.11891412734985352
8 4.53515625 0.01224285364151001
9 4.529296875 0.053596146404743195
10 4.5322265625 0.020741849206387997
11 4.53369140625 0.0042658079182729125
12 4.534423828125 0.003984444148954935
13 4.5340576171875 0.0001417014154867502
14 4.53424072265625 0.0019211164656098845
15 4.534149169921875 0.0008896438020826736
16 4.5341033935546875 0.0003739552628445608
17 4.534080505371094 0.0001161229410939768
18 4.534069061279297 1.2790232830184323e-05
19 4.534074783325195 5.166610522167048e-05
20 4.534071922302246 1.9437873959304852e-05
21 4.5340704917907715 3.3238050036743516e-06
22 4.534069776535034 4.7332178070291775e-06
23 4.534070134162903 7.047073751209609e-07
1 4.5 0.375
2 5.25 12.328125
3 4.875 4.763671875
4 4.6875 1.910888671875
5 4.59375 0.699554443359375
6 4.546875 0.14548873901367188
7 4.5234375 0.11891412734985352
8 4.53515625 
```