## Exercise 04.1 (simple function)

Write a function called `is_odd` which takes an integer as an argument and returns `True` if the argument is odd, and otherwise returns `False`. Test your function for several values.

<!-- grade_id: cell-f07e8ac1f9838913, points: 0, grade: false, solution: true -->
```python
def is_odd(x):
    if x % 2 == 0:
        return(False)
    else:
        return(True)
```

<!-- grade_id: cell-8a4ca795d9fa96dd, points: 0, grade: true, solution: false -->
```python
assert is_odd(0) == False
assert is_odd(101) == True
assert is_odd(982) == False 
assert is_odd(-5) == True
assert is_odd(-8) == False
```

## Exercise 04.2 (functions and default arguments)

Write a single function named `magnitude` that takes each component of a vector of length 2 or 3 and returns the magnitude.
Use default arguments to handle vectors of length 2 or 3 with the same code. Test your function for correctness against hand calculations for a selection of values.

<!-- grade_id: cell-693a09b45c6bb532, points: 0, grade: false, solution: true -->
```python
def magnitude(*args):
    import math
    sum = 0
    for i in args:
        sum += i**2
    return(math.sqrt(sum))
```

<!-- grade_id: cell-384a85ba97b226d7, points: 0, grade: true, solution: false -->
```python
assert round(magnitude(3, 4) - 5.0, 10) == 0.0
assert round(magnitude(4, 3) - 5.0, 10) == 0.0
assert round(magnitude(4, 3, 0.0)- 5.0, 10) == 0.0
assert round(magnitude(4, 0.0, 3.0) - 5.0, 10) == 0.0
assert round(magnitude(3, 4, 4) - 6.403124237, 8) == 0.0
```

## Exercise 04.3 (functions)

Given the coordinates of the vertices of a triangle, $(x_0, y_0)$, $(x_1, y_1)$ and $(x_2, y_2)$, the area $A$ of the triangle is given by:
$$
A = \left| \frac{x_0(y_1  - y_2) + x_1(y_2 - y_0) + x_2(y_0 - y_1)}{2} \right|
$$
Write a function named `area` that computes the area of a triangle given the coordinates of the vertices.
The order of the function arguments must be (`x0, y0, x1, y1, x2, y2)`.

Test the output of your function against some known solutions.

<!-- grade_id: cell-ed02439d4509b2c1, points: 0, grade: false, solution: true -->
```python
def area(x0,y0,x1,y1,x2,y2):
    area = abs((x0*(y1-y2) + x1*(y2-y0) + x2*(y0-y1))/2)
    return(area)
```

<!-- grade_id: cell-64f8b6708595aed3, points: 0, grade: true, solution: false -->
```python
x0, y0 = 0.0, 0.0
x1, y1 = 0.0, 2.0
x2, y2 = 3.0, 0.0
A = area(x0, y0, x1, y1, x2, y2)
assert round(A - 3.0, 10) == 0.0
```

## Exercise 04.4 (recursion)

The factorial of a non-negative integer $n$ is expressed recursively by:
$$
n! = 
\begin{cases}
1 & n = 0 \\
(n - 1)! \,n & n > 0
\end{cases}
$$

Develop a function named `factorial` for computing the factorial using recursion.
Test your function against the `math.factorial` function, e.g.

```python
import math
print("Reference factorial:", math.factorial(5))
```

Output:
```
Reference factorial: 120
```

<!-- grade_id: cell-9591c1fc1c49c167, points: 0, grade: false, solution: true -->
```python
def factorial(n):
    if n == 0:
        return(1)
    else:
        fact = 1
        for i in range(1,n+1):
            # print(f'i = {i}')
            fact = fact*i
            # print(f'Factorial = {fact}')
        return(fact)
    
print("Factorial of 5:", factorial(5))

import math
print("Reference value of factorial of 5:", math.factorial(5))
```

Output:
```
Factorial of 5: 120
Reference value of factorial of 5: 120
```

<!-- grade_id: cell-ffe589e9fb2f1419, points: 0, grade: true, solution: false -->
```python
assert factorial(0) == 1
assert factorial(1) == 1
assert factorial(2) == 2
assert factorial(5) == 120

import math
assert factorial(32) == math.factorial(32)
```

## Exercise 04.5 (functions and passing functions as arguments)

Restructure your program from the bisection problem in Exercise 02 to 

- Use a Python function to evaluate the mathematical function $f$ that we want to find the root of; 

and then

- Encapsulate the bisection algorithm inside a Python function, which takes as arguments:
  - the function we want to find the roots of
  - the points $x_{0}$ and $x_{1}$ between which we want to search for a root
  - the tolerance for exiting the bisection algorithm (exit when $|f(x)| < \text{tol}$)
  - maximum number of iterations (the algorithm should exit once this limit is reached)

For the first step, use a function for evaluating $f$, e.g.:
```python
def f(x):
    # Put body of the function f(x) here, and return the function value
```           
For the second step, encapsulate the bisection algorithm in a function:
```python
def compute_root(f, x0, x1, tol, max_it):
    # Implement bisection algorithm here, and return when tolerance is satisfied or
    # number of iterations exceeds max_it

    # Return the approximate root, value of f(x) and the number of iterations
    return x, f, num_it

# Compute approximate root of the function f
x, f_x, num_it = compute_root(f, x0=3, x1=6, tol=1.0e-6, max_it=1000)
```

Try testing your program for a different function. A quadratic function, whose roots you can find analytically, would be a good test case.


### Optional extension 

Use recursion to write a `compute_root` function that *does not* require a `for` or `while` loop.

### Solution

Define the function for computing $f(x)$:

```python
def my_f(x):
    "Evaluate polynomial function"
    return x**3 - 6*x**2 + 4*x + 12
```

Create the function that performs the bisection:

<!-- grade_id: cell-622edcb129dcc7f2, points: 0, grade: false, solution: true -->
```python
def compute_root(f, x0, x1, tol, max_it):
    
    "Compute roots of a function using bisection"
    error = tol + 1.0

    # Iterate until tolerance is met
    it = 0
    while error > tol:
        
        # Iteration counter
        it += 1
        
        # Guard against an infinite loop
        if it > max_it:
            print("Oops, iteration count is very large. Breaking out of while loop.")
            break
        
        # Compute midpoint
        x_mid = (x0 + x1)/2

        # Evaluate function at left end-point and at midpoint
        f0 = x0**3 - 6*x0**2 + 4*x0 + 12
        f = x_mid**3 - 6*x_mid**2 + 4*x_mid + 12
        f_mid = f
        
        if f0 * f < 0:
            x1 = x_mid
        else:
            x0 = x_mid
        
        # Computating error
        error = abs(f)
        
        print(it, x_mid, error)
     
    return x_mid, f_mid, it
```

<!-- grade_id: cell-cb5a5cdf1d790edb, points: 0, grade: true, solution: false -->
```python
x, f, num_it = compute_root(my_f, x0=3, x1=6, tol=1.0e-6, max_it=1000)
assert round(x - 4.534070134162903, 10) == 0.0
```

Output:
```
1 4.5 0.375
2 5.25 12.328125
3 4.875 4.763671875
4 4.6875 1.910888671875
5 4.59375 0.699554443359375
6 4.546875 0.14548873901367188
7 4.5234375 0.11891412734985352
8 4.53515625 0.01224285364151001
9 4.529296875 0.053596146404743195
10 4.5322265625 0.020741849206387997
11 4.53369140625 0.0042658079182729125
12 4.534423828125 0.003984444148954935
13 4.5340576171875 0.0001417014154867502
14 4.53424072265625 0.0019211164656098845
15 4.534149169921875 0.0008896438020826736
16 4.5341033935546875 0.0003739552628445608
17 4.534080505371094 0.0001161229410939768
18 4.534069061279297 1.2790232830184323e-05
19 4.534074783325195 5.166610522167048e-05
20 4.534071922302246 1.9437873959304852e-05
21 4.5340704917907715 3.3238050036743516e-06
22 4.534069776535034 4.7332178070291775e-06
23 4.534070134162903 7.047073751209609e-07
```

#### Optional extension: using recursion:

<!-- grade_id: cell-c944bb2229317d4d, points: 0, grade: false, solution: true -->
```python
def compute_root(f, x0, x1, tol, max_it, it_count=0):
    # Iteration counter
    it_count += 1
    
    # Compute midpoint
    x_mid = (x0 + x1)/2

    # Evaluate function at left end-point and at midpoint
    f0 = x0**3 - 6*x0**2 + 4*x0 + 12
    f = x_mid**3 - 6*x_mid**2 + 4*x_mid + 12
    
    if f0 * f < 0:
        x1 = x_mid
    else:
        x0 = x_mid
    
    # Computating error
    error = abs(f)
    
    print(it_count, error)
    
    if error < tol or it_count > max_it:
        # Call compute_root recursively
        return (x0,f,it_count)
    else:
        return compute_root(f, x0, x1, tol, max_it, it_count)
        
```

<!-- grade_id: cell-9ebe2f13b2899a8f, points: 0, grade: true, solution: false -->
```python
x, f, num_it = compute_root(my_f, x0=3, x1=6, tol=1.0e-6, max_it=1000)
assert round(x - 4.534070134162903, 10) == 0.0
```

Output:
```
1 0.375
2 12.328125
3 4.763671875
4 1.910888671875
5 0.699554443359375
6 0.14548873901367188
7 0.11891412734985352
8 0.01224285364151001
9 0.053596146404743195
10 0.020741849206387997
11 0.0042658079182729125
12 0.003984444148954935
13 0.0001417014154867502
14 0.0019211164656098845
15 0.0008896438020826736
16 0.0003739552628445608
17 0.0001161229410939768
18 1.2790232830184323e-05
19 5.166610522167048e-05
20 1.9437873959304852e-05
21 3.3238050036743516e-06
22 4.7332178070291775e-06
23 7.047073751209609e-07
```