## Exercise 06.1 (selecting and passing data structures)

The task in Exercise 04 for computing the area of a triangle involved a function with six arguments ($x$ and $y$ components of each vertex). With six arguments, the likelihood of a user passing arguments in the wrong order is high. 

Use an appropriate data structure, e.g. a `list`, `tuple`, `dict`, etc,  to develop a new version of the function with a simpler interface (the interface is the arguments that are passed to the function). Add appropriate checks inside your function to validate the input data.

<!-- grade_id: cell-50c4566d7bcbffaf, points: 0, grade: false, solution: true -->
```python
def area(x0,y0,x1,y1,x2,y2):
    area = abs((x0*(y1-y2) + x1*(y2-y0) + x2*(y0-y1))/2)
    return(area)
```

## Exercise 06.2 (selecting data structures)

For a simple (non-intersecting) polygon with $n$ vertices, $(x_0, y_0)$, $(x_1, y_1)$, . . , $(x_{n-1}, y_{n-1})$, the area $A$ is given by
$$
A = \left| \frac{1}{2} \sum_{i=0}^{n-1} \left(x_{i} y _{i+1} - x_{i+1} y_{i} \right) \right|
$$
and where $(x_n, y_n) = (x_0, y_0)$. The vertices should be ordered as you move around the polygon.

Write a function that computes the area of a simple polygon with an arbitrary number of vertices. Test your function for some simple shapes. Pay careful attention to the range of any loops.

<!-- grade_id: cell-4b5a619405787746, points: 0, grade: false, solution: true -->
```python
def polygon_area(*args):
    area = 0
    for i in range(0,len(args),2):
        # x_i and y_i
        xi = args[i]
        yi = args[i+1]
        
        # x_i+1 and y_i+1
        if (i+3) > len(args):
            xip1 = args[i]
            yip1 = args[i+1]
        else:
            xip1 = args[i+2]
            yip1 = args[i+3]
        
        # calculating sides
        first = (xi*yip1)
        second = (xip1*yi)
        area += (1/2)*(first - second)
    return(abs(area))

x0, y0 = 0.0, 0.0
x1, y1 = 0.0, 2.0
x2, y2 = 3.0, 2.0
x3, y3 = 3.0, 0.0
A = polygon_area(x0, y0, x1, y1, x2, y2)
# print(A)
assert round(A - 3.0, 10) == 0.0

A = polygon_area(x0, y0, x1, y1, x2, y2, x3, y3)
# print(A)
assert round(A - 6.0, 10) == 0.0


```

## Exercise 06.3 (indexing)

Write a function that uses list indexing to add two vectors of arbitrary length, and returns the new vector. Include a check that the vector sizes match, and print a warning message if there is a size mismatch. The more error information you provide, the easier it would be for someone using your function to debug their code.

Add some tests of your code.

#### Hint: You can create a list of zeros of length `n` by

    z = [0]*n
    
#### Optional (advanced) 

Try writing a one-line version of this operation using list comprehension and the built-in function [`zip`](https://docs.python.org/3/library/functions.html#zip).

<!-- grade_id: cell-47daca03f7a96b62, points: 0, grade: false, solution: true -->
```python
def sum_vector(x, y):
    "Return sum of two vectors"
    sum = []
    for i,j in zip(x,y):
        sum.append(i+j)
    return(sum)
```

<!-- grade_id: cell-95f5819810d2b2c7, points: 0, grade: true, solution: false -->
```python
a = [0, 4.3, -5, 7]
b = [-2, 7, -15, 1]
c = sum_vector(a, b)
assert c == [-2, 11.3, -20, 8]
```

### Extension: list comprehension

<!-- grade_id: cell-c182b3f0d8952692, points: 0, grade: true, solution: true -->
```python
def sum_vector(x, y):
    "Return sum of two vectors"
    sum = [i+j for i,j in zip(x,y)]
    return(sum)
```

## Exercise 06.4 (dictionaries)

Create a dictionary that maps college names (the key) to college abbreviations for at least 5 colleges
(you can find abbreviations at https://en.wikipedia.org/wiki/Colleges_of_the_University_of_Cambridge#Colleges).
From the dictionary, produce and print

1. A dictionary from college abbreviation to name; and
1. A list of college abbreviations sorted into alphabetical order.

*Optional extension:* Create a dictionary that maps college names (the key) to dictionaries of:

- College abbreviation
- Year of foundation 
- Total number students
 
for at least 5 colleges. Take the data from https://en.wikipedia.org/wiki/Colleges_of_the_University_of_Cambridge#Colleges. Using this dictionary, 

1. Find the college with the greatest number of students and print the abbreviation; and 
2. Find the oldest college, and print the number of students and the abbreviation for this college.

<!-- grade_id: cell-bc271dc3885e19c5, points: 0, grade: false, solution: true -->
```python
colleges = {'Downing':{'Abbreviation':'DOW','YOF':1800,'Students':845},
            'Darwin':{'Abbreviation':'DAR','YOF':1964,'Students':755},
            'Jesus':{'Abbreviation':'JE','YOF':1496,'Students':924},
            'Magdalene':{'Abbreviation':'M','YOF':1428,'Students':572},
            'Robinson':{'Abbreviation':'R','YOF':1977,'Students':664}}
```

#### Optional extension

<!-- grade_id: cell-7040c4053c72ce09, points: 0, grade: false, solution: true -->
```python
top_college = ""
max_stud = 0
for d in colleges:
    temp = colleges[d]['Students']
    if temp > max_stud:
        max_stud = temp
        top_college = d

print(max_stud)
print(top_college)

# The bulk is done... skipping the easy part.

```

Output:
```
924
Jesus
```