Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

---

# Advanced concepts

This lab provides exercises for decorators and iterators. Complete the assignments and return this notebook `advanced-2020-05-13.ipynb` to Canvas

## Context managers

Context managers are statements with a `with` block and are implemented by classes with `__enter__` and `__exit__` methods. The official documentation is at https://docs.python.org/3/library/stdtypes.html#context-manager-types

### Assignment 1

Given the color switch escape codes for terminals

    >>> GREEN = "\033[32m"
    >>> RED = "\033[31m"
    >>> RESET = "\033[00m"

Having the property illustrated by the following block:

```python
GREEN = "\033[32m"
RED = "\033[31m"
RESET = "\033[00m"

print (GREEN + "yes", RED + "no", RESET + 'maybe')
```

Output:
```
[32myes [31mno [00mmaybe
```

Complete a context manager that changes the color of the output such that

<code>
>>> with Color(RED):
...     print('hello')    
<span style="color:red">hello</span>
>>> with Color(GREEN):
...     print('hello')    
<span style="color:green">hello</span>
>>>
</code>


<!-- grade_id: cell-e1388e8c145cd2ca, points: 0, grade: false, solution: true -->
```python
class Color:
    GREEN = "\033[32m"
    RED = "\033[31m"
    RESET = "\033[00m"

    def __init__(self, color):
        self.color = color
        
    def __enter__(self): # CALLED AFTER "WITH"
        print(self.color)
        
    def __exit__(self, *args): # CALLED AT THE "END"
        print(RESET, end='')
```

<!-- grade_id: cell-73773482ff268d10, points: 1, grade: true, solution: false -->
```python
# Verify that you get color output

with Color(RED):
    print('Red day')
    
with Color(GREEN):
    print('Green forest')
    
```

Output:
```
[31m
Red day
[00m[32m
Green forest
[00m
```

## Decorators

A decorator is a function that modifies the behaviour of another function.

The syntax

~~~
@decorator1
@decorator2
def f():
   ...
~~~

is so called syntactic sugar for

~~~
def f():
    ...
~~~

f = decorator1(decorator2(f))

The first form is convenient for our own code, but the second form must be used when we do not have the source at hand, e.g. for built-in functions

### Assignment 2

Repeat the functionality of Assignment1 for decorators such that a decorated function gets color output:

<code>
@print_in_red
def hello():
    print("Hello world!")
</code>

<code>
>>> hello()
<span style="color:red">Hello world!</span>
</code>
>>>

<!-- grade_id: cell-1282ae349daeef71, points: 0, grade: false, solution: true -->
```python
def print_in_red(f):
    RED = "\033[31m"
    RESET = "\033[00m"

    def wrapper(*args, **kwargs):
        print(RED)
        wrap_return = f(*args, **kwargs)
        print(RESET)
        return(wrap_return)
    return(wrapper)
```

<!-- grade_id: cell-19eb7c83720e263e, points: 1, grade: true, solution: false -->
```python
# Verify that you get color output

@print_in_red
def hello():
    print("Hello world!")
     
hello()

print('Test if color is reseted')
```

Output:
```
[31m
Hello world!
[00m
Test if color is reseted
```

## Iterators/generators

### Assignment 3

The `datetime` module contains a number of helper function to handle dates and times. Here we will use the `date` and `timedelta` objects. 

* `date` objects refers to date information with year/month/day and more
* `timedelta` objects are differences between two dates (or date-times). They can be added to date objects to obtain new dates in some future.

Example of how they work is

```python
import datetime
today = datetime.date.today()
today
```

Output:
```
datetime.date(2022, 3, 7)
```

```python
week = datetime.timedelta(days=7)
today + week
```

Output:
```
datetime.date(2022, 3, 14)
```

Write a generator that returns the dates a number of weeks from now, in string format using `isoformat` method

<!-- grade_id: cell-03b5401cc183bf8e, points: 0, grade: false, solution: true -->
```python
def date_weeks_ahead(starting_date, number_of_weeks):
    week = datetime.timedelta(days=7)
    start = starting_date
    dates = []
    for i in range(1,number_of_weeks+1):
        dates.append(str(start + week*i))
    return(dates)
```

<!-- grade_id: cell-b42fc950abbaa239, points: 1, grade: true, solution: false -->
```python
today = datetime.date(2022, 3, 4)
actual = list(date_weeks_ahead(today, 4))
expected = ['2022-03-11', '2022-03-18', '2022-03-25', '2022-04-01']
assert actual == expected, f'{actual} != {expected}'
```

# PAST EXAM QUESTIONS

### The `partial()` function

The `partial()` is used for partial function application which `“freezes”` some portion of a function’s arguments and/or keywords resulting in a new object
with a simplified signature. For example, `partial()` can be used to create a callable that behaves like the `int()` function where the base argument
defaults to two:

    >>> from functools import partial
    >>> basetwo = partial(int, base=2)
    >>> basetwo.__doc__ = 'Convert base 2 string to an int.'
    >>> basetwo('10010')
    18

Make an analogy of this for the print function such that objects are printed on separate lines.

Hint: use the `sep` keyword argument

Solution:

    >>> lprint = partial(print, sep='\n')
    
    >>> lprint('a', 'b', 'c')
    a
    b
    c

#### Return a `zip object` whose `.next()` method returns a tuple

The zip documentation contains

`class zip(object) zip(iter1 [,iter2 [...]]) --> zip object`

Return a `zip object` whose `.next()` method returns a `tuple` where the `i-th element` comes from the `i-th iterable argument.` The `.next()` method
continues until the shortest iterable in the argument sequence is exhausted and then it raises `StopIteration`.

Solution: 

    >>> l1 = [1, 2, 3]
    >>> l2 = [4, 5]
        >>> for z in zip(l1, l2):
        >>> print(z)

    (1, 4)
    (2, 5)

### The `map` function

The map function has the following documentation

`class map(object) map(func, iterables) --> map object*`

`Make an iterator that computes the function using arguments from each of the iterables. Stops when the shortest iterable is exhausted.`

What is the output of the following?

    l1 = [1, 2, 3]
    l2 = [4, 5]
    def f(x, y): return x + y
    for s in map(f, l1, l2):
    print(s)

Solution:

    5
    7

### The `colorama module` --> color output in a terminal

The colorama module in Python can be used to give color output in a terminal e.g.

    >>> from colorama import Fore, Style
    >>> print(Fore.RED + 'some red text' + Style.RESET_ALL)
    some red text

Use this to design a decorator such that all output from a decorated function is in red, such that:

    >>> @red
        def hello():
        ... print("Hello world!")
    >>> hello()
    Hello world! (in red)

Solution:

```python
def red(f):
    def wrap(*args, **kwargs):
        print(Fore.RED, end='') # '' Removes empty printed lines
        f(*args, **kwargs)
        print(Style.RESET_ALL)
    return wrap
```

### Identity Decorators

How do you write and apply an identity decorator, which does not modify
the function it is applied to?

    def identity(func):
        return func

### Datatypes of `*args` and `**kwargs`

A general function definition has starred arguments:

    >>> def f(*args, **kwargs):
            print(type(args).__name__)
            print(type(kwargs).__name__)

What is the data type of args and kwargs respectively?

    >>> f()
    tuple = "args"
    dict = "kwargs"

### Decorator that runs a function twice

Such that:

    @dotwice
    def hello():
        print("Hello")
        
    >>> hello()
    Hello
    Hello

```python
def dotwice(func):
    def wrapper(*args, **kwargs):
        func(*args, **kwargs)
        func(*args, **kwargs)
    return wrapper
```

### Accessing a decorator function without the `@decorator` 

The decorator syntax is handy when you apply it to your own code, e.g.

    @timer
    def myfunc():
    ...

Another syntax must be use if you want to apply it to a library function without direct access to the code. 

What would you write to apply the math.sqrt function?

    math.sqrt = timer(math.sqrt)