Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

---

# Basic Python syntax


## Exercise 1
Write a function 'mymax' that accepts two numbers as parameters and that prints the larger number to the screen. For this elementary exercise, do not use the built-in function 'max()' - just revert to elementary expressions and statements.

<!-- grade_id: cell-ee8c56c7cc588b03, points: 0, grade: false, solution: true -->
```python
def mymax(x, y):
    """
    Returns the maximum of x and y
    """
    if x > y:
        return(x)
    elif y > x:
        return(y)
    else:
        return(x)
        print('x = y')

#raise NotImplementedError()
```

<!-- grade_id: cell-304328f9c7d4c051, points: 1, grade: true, solution: false -->
```python
assert mymax(2, 6) == 6, f'mymax(2, 6) -> {mymax(2, 6)} != 6'
assert mymax(5, 3) == 5, f'mymax(5, 3) -> {mymax(5, 3)} != 5'
```

## Exercise 2

Write a function to sum the elements of a sequence with a for loop. You can assume that it consists of numeric values.

<!-- grade_id: cell-b607688727ab57f1, points: 0, grade: false, solution: true -->
```python
def mysum(sequence):
    """
    Returns the sum of the elements in sequence
    """
    s = 0
    for x in sequence:
        s += x
    
    return s

#raise NotImplementedError()
```

<!-- grade_id: cell-24ac435261756312, points: 1, grade: true, solution: false -->
```python
assert mysum([1, 2]) == 3, f'mysum([1, 2]) -> {mysum([1, 2])} != 3'
assert mysum((1, 1/2, 1/4)) == 1.75, f'mysum((1, 1/2, 1/4)) -> {mysum((1, 1/2, 1/4))} != 1.75'
```

## Exercise 3

Variables of `str` type has a method `join` defined

Variables of `str` type has a method `join` defined such that when a sequence (e.g. `list`, `tuple`) of strings is given as an argument, they are concatenated to a single string with the variable as a separator. E.g.

    >>> separator = '-'
    >>> separator.join(['father', 'in', 'law'])
    'father-in-law'

Use the `join` to generate a histogram with stars such that 

    histogram([3, 4, 1]) 
    
returns

    ***
    ****
    *

<!-- grade_id: cell-631ae09a389a66d4, points: 0, grade: false, solution: true -->
```python
def histogram(sequence):
    """
    Returns histogram in the form of characters
    >>> histogram([3, 4, 1])
    ***
    ****
    *
    """

    histo = []
    separator1 = ""
    separator2 = "\n"
    for x in sequence:
        hist = []
        for i in range(x):
            hist.append("*")

        histo.append(separator1.join(hist))
    return separator2.join(histo)
            
#raise NotImplementedError()
    

print(f'Example: the histogram of [3, 4, 1] is \n{histogram([3, 4, 1])}')
```

<!-- grade_id: cell-901ea19debe3ee4c, points: 1, grade: true, solution: false -->
```python
assert histogram([]) == "", f'"{histogram([])}" != "" '
assert histogram([1, 2]) == "*\n**", f'"{histogram([1, 2])}" != \n"*\n**"'
assert histogram([1, 2, 3]) == "*\n**\n***", f'"{histogram([1, 2, 3])}" != \n "*\n**\n***"'
```

## Exercise 4

Compute how many times a character is present in a string. Use a dictionary to represent the string. 

<!-- grade_id: cell-94b1ecc182583f9b, points: 0, grade: false, solution: true -->
```python
def character_count(text):
    """
    Returns character count of text
    Input: text of type str
    Output: dict that maps a character to a number count
    """
    d = {}
    for i in text:
        c = text.count(i)
        d[i] = c
    return(d)
    #raise NotImplementedError()

print(f"Example: Character count of 'ABBA': {character_count('ABBA')}")
```

<!-- grade_id: cell-2a787760db753909, points: 1, grade: true, solution: false -->
```python
calculated = character_count("ABBA")
expected = {"A": 2, "B": 2}
assert calculated == expected, f'{calculated} != {expected}'

calculated = character_count("CCGGAAGAGCTTACTTAG")
expected = {"A": 5, "C": 4, "G": 5, "T": 4}
assert calculated == expected, f'{calculated} != {expected}'
```

## Exercise 5

A simple encryption scheme known as Caesarian roll

A simple encryption scheme known as Caesarian roll, is obtained by shifting characters a fixed amount. In Python we can obtain the order of a character with the `ord` function

    >>> ord('A')
    65
    
and there reverse operation

    >>> chr(65)
    'A' 
    
Use this in a function that takes a string and returns the encrypted representation

<!-- grade_id: cell-95b86ffef93462ec, points: 0, grade: false, solution: true -->
```python
def caesar_shift(text, n):
    """
    Returns an encrypted string by shifting input characters by n
    """
    t = []
    sep = ""
    for i in text:
        o = ord(i)
        o_new = chr(o + n)
        t.append(o_new)
    t = sep.join(t)
    return(t)
    
    #raise NotImplementedError()

print(f"Example: Caesarian roll of 'ABBA' and shift = 2: {caesar_shift('ABBA',2)}")
```

<!-- grade_id: cell-4403018cd26de943, points: 1, grade: true, solution: false -->
```python
assert caesar_shift('a', 0) == 'a'
assert caesar_shift('ABBA', 2) == 'CDDC'
assert caesar_shift('CDDC', -2) == 'ABBA'
```

Optional:

Historically, this encryption allegedly used  by the Roman Emperor Caesar (thereby the name). Eventually the encryption was crackedby realizing that some characters are more common than other in a given language. Use the functions defined in previous exercises to

* identify the most common character in the crypto text
* in the text the second common character is 'e' - what is the shift?
* decrypt the crypto_text with `caesar_shift` using a negative shift

```python
crypto_text = "{v'il3'vy'uv{'{v'ilA'{oh{'pz'{ol'x|lz{pvuA\x11~ol{oly'.{pz'uvisly'pu'{ol'tpuk'{v'z|mmly\x11{ol'zspunz'huk'hyyv~z'vm'v|{yhnlv|z'mvy{|ul3\x11vy'{v'{hrl'hytz'hnhpuz{'h'zlh'vm'{yv|islz3\x11huk'i\x80'vwwvzpun'luk'{oltF'{v'kplA'{v'zsllwB\x11uv'tvylB'huk'i\x80'h'zsllw'{v'zh\x80'~l'luk\x11{ol'olhy{4hjol'huk'{ol'{ov|zhuk'uh{|yhs'zovjrz\x11{oh{'mslzo'pz'olpy'{v3'.{pz'h'jvuz|tth{pvu\x11kl}v|{s\x80'{v'il'~pzo.k5'{v'kpl3'{v'zsllwB\x11{v'zsllwA'wlyjohujl'{v'kylhtA'h\x803'{olyl.z'{ol'y|iB\x11mvy'pu'{oh{'zsllw'vm'klh{o'~oh{'kylhtz'th\x80'jvtl\x11~olu'~l'oh}l'zo|mmslk'vmm'{opz'tvy{hs'jvps3\x11t|z{'np}l'|z'wh|zlA'{olyl.z'{ol'ylzwlj{\x11{oh{'thrlz'jhshtp{\x80'vm'zv'svun'spmlB\x11mvy'~ov'~v|sk'ilhy'{ol'~opwz'huk'zjvyuz'vm'{ptl3\x11{ol'vwwylzzvy.z'~yvun3'{ol'wyv|k'thu.z'jvu{|tls\x803\x11{ol'whunz'vm'klzwpzlk'sv}l3'{ol'sh~.z'klsh\x803\x11{ol'puzvslujl'vm'vmmpjl'huk'{ol'zw|yuz\x11{oh{'wh{plu{'tlyp{'vm'{ol'|u~vy{o\x80'{hrlz3\x11~olu'ol'optzlsm'tpno{'opz'x|pl{|z'thrl\x11~p{o'h'ihyl'ivkrpuF'~ov'~v|sk'mhyklsz'ilhy3\x11{v'ny|u{'huk'z~lh{'|ukly'h'~lhy\x80'spml3\x11i|{'{oh{'{ol'kylhk'vm'zvtl{opun'hm{ly'klh{o3\x11{ol'|ukpzjv}ly.k'jv|u{y\x80'myvt'~ovzl'iv|yu\x11uv'{yh}lssly'yl{|yuz3'w|\x81\x81slz'{ol'~pss\x11huk'thrlz'|z'yh{oly'ilhy'{ovzl'pssz'~l'oh}l\x11{ohu'ms\x80'{v'v{olyz'{oh{'~l'ruv~'uv{'vmF\x11{o|z'jvuzjplujl'kvlz'thrl'jv~hykz'vm'|z'hssB\x11huk'{o|z'{ol'uh{p}l'o|l'vm'ylzvs|{pvu\x11pz'zpjrsplk'v.ly'~p{o'{ol'whsl'jhz{'vm'{ov|no{3\x11huk'lu{lywypzlz'vm'nylh{'wp{o'huk'tvtlu{\x11~p{o'{opz'ylnhyk'{olpy'j|yylu{z'{|yu'h~y\x803\x11huk'svzl'{ol'uhtl'vm'hj{pvu5\x11"
```

<!-- grade_id: cell-440cb014209f1873, points: 0, grade: true, solution: true -->
```python
crypto_text = "{v'il3'vy'uv{'{v'ilA'{oh{'pz'{ol'x|lz{pvuA\x11~ol{oly'.{pz'uvisly'pu'{ol'tpuk'{v'z|mmly\x11{ol'zspunz'huk'hyyv~z'vm'v|{yhnlv|z'mvy{|ul3\x11vy'{v'{hrl'hytz'hnhpuz{'h'zlh'vm'{yv|islz3\x11huk'i\x80'vwwvzpun'luk'{oltF'{v'kplA'{v'zsllwB\x11uv'tvylB'huk'i\x80'h'zsllw'{v'zh\x80'~l'luk\x11{ol'olhy{4hjol'huk'{ol'{ov|zhuk'uh{|yhs'zovjrz\x11{oh{'mslzo'pz'olpy'{v3'.{pz'h'jvuz|tth{pvu\x11kl}v|{s\x80'{v'il'~pzo.k5'{v'kpl3'{v'zsllwB\x11{v'zsllwA'wlyjohujl'{v'kylhtA'h\x803'{olyl.z'{ol'y|iB\x11mvy'pu'{oh{'zsllw'vm'klh{o'~oh{'kylhtz'th\x80'jvtl\x11~olu'~l'oh}l'zo|mmslk'vmm'{opz'tvy{hs'jvps3\x11t|z{'np}l'|z'wh|zlA'{olyl.z'{ol'ylzwlj{\x11{oh{'thrlz'jhshtp{\x80'vm'zv'svun'spmlB\x11mvy'~ov'~v|sk'ilhy'{ol'~opwz'huk'zjvyuz'vm'{ptl3\x11{ol'vwwylzzvy.z'~yvun3'{ol'wyv|k'thu.z'jvu{|tls\x803\x11{ol'whunz'vm'klzwpzlk'sv}l3'{ol'sh~.z'klsh\x803\x11{ol'puzvslujl'vm'vmmpjl'huk'{ol'zw|yuz\x11{oh{'wh{plu{'tlyp{'vm'{ol'|u~vy{o\x80'{hrlz3\x11~olu'ol'optzlsm'tpno{'opz'x|pl{|z'thrl\x11~p{o'h'ihyl'ivkrpuF'~ov'~v|sk'mhyklsz'ilhy3\x11{v'ny|u{'huk'z~lh{'|ukly'h'~lhy\x80'spml3\x11i|{'{oh{'{ol'kylhk'vm'zvtl{opun'hm{ly'klh{o3\x11{ol'|ukpzjv}ly.k'jv|u{y\x80'myvt'~ovzl'iv|yu\x11uv'{yh}lssly'yl{|yuz3'w|\x81\x81slz'{ol'~pss\x11huk'thrlz'|z'yh{oly'ilhy'{ovzl'pssz'~l'oh}l\x11{ohu'ms\x80'{v'v{olyz'{oh{'~l'ruv~'uv{'vmF\x11{o|z'jvuzjplujl'kvlz'thrl'jv~hykz'vm'|z'hssB\x11huk'{o|z'{ol'uh{p}l'o|l'vm'ylzvs|{pvu\x11pz'zpjrsplk'v.ly'~p{o'{ol'whsl'jhz{'vm'{ov|no{3\x11huk'lu{lywypzlz'vm'nylh{'wp{o'huk'tvtlu{\x11~p{o'{opz'ylnhyk'{olpy'j|yylu{z'{|yu'h~y\x803\x11huk'svzl'{ol'uhtl'vm'hj{pvu5\x11"
#print("Crypto text:\n", crypto_text)

# CREATING DICTIONARY AND SORTED DICTIONARY OF CRYPTO TEXT
def character_count(text):
    d = {}
    for i in text:
        c = text.count(i)
        d[i] = c
    dsort = sorted(d.items(), key=lambda x: x[1], reverse=True)
    return(d,dsort)

# FINDING CAESARIAN SHIFT
def find_shift(dsort):
    seckey = ord(dsort[1][0])
    e_ord = ord("e")
    shift = e_ord - seckey
    return(shift)

# FINDING MOST COMMON CHARACTER IN DECRYPTED TEXT
def find_common(dsort):
    comkey = ord(dsort[0][0])
    comkeyshift = chr(comkey + roll)
    return(comkeyshift)

# DECRYPTING TEXT
def caesar_shift(crypto_text,roll):
    decrypt = []
    for i in crypto_text:
        o = chr(ord(i) + roll)
        decrypt.append(o)
    decrypt = "".join(decrypt)
    return decrypt

# GENERATING OUTPUT
d_out, dsort_out = character_count(crypto_text)
roll = find_shift(dsort_out)
com_char = find_common(dsort_out)

print(f"Dictionary: {character_count(crypto_text)[0]} \n")
print(f"Sorted list: {character_count(crypto_text)[1]} \n")
print(f"Shift: '{find_shift(dsort_out)}' \n")
print(f"Most common character: '{find_common(dsort_out)}' \n")

print("\nClear text:\n", caesar_shift(crypto_text, roll))

```

# PAST EXAM QUESTIONS:

### Write a function definition that (input = none, list, zero, etc)

- takes one optional parameter with default value None
- with no parameter it returns -999
- the parameter can be assumed to be a sequence-like type (e.g. list, tuple)
- the function returns the sum of the parameter elements if provided
- an empty sequence gives zero

```python
def sum_elements(seq=None):
    if seq is None:
        return -999
    else:
        return sum(seq)
```

Such that:

    >>> sum_elements() = -999
    >>> sum_elements([]) = 0
    >>> sum_elements((1, 2)) = 3
    >>> sum_elements(range(4)) = 6

### Loop over a dictionary

List a way of looping over a dictionary (there are a few) so that e.g.{ "a": 1, "b": 2, "c": 3} is displayed on the screen as

    >>> a->1
    >>> b->2
    >>> c->3

```python
dictex = {"a": 1, "b": 2, "c": 3}

def dictloop(d):
    for k, v in d.items():
        print(f"{k}->{v}")

dictloop(dictex)
```

### Creating dictionary with gender as key and full names as values

A function takes a CSV-string with name/gender data as input and returns a dictionary with gender as key and a list of full names as value.
e.g. with

    >>> sample_data = """first_name,last_name,gender
    ... Quintina,Firle,F
    ... Jesse,Nunson,M
    ... Lena,Stockley,F
    ... Margaux,Sirr,F
    ... Taylor,Alishoner,M"""

Reorder and indent correctly the source lines:

    >>> names = {"F": [], "M": []}
    ... return names

And then

    >>> def group_by_gender(data):
    ... first, last, gender = line.split(",")
    ... lines = data.split("\n")
    ... names[gender].append(f"{first} {last}")
    ... for line in lines[1:]:

```python
def group_by_gender(data):
    names = {"F": [], "M": []}
    lines = data.split("\n")
    for line in lines[1:]:
        first, last, gender = line.split(",")
        names[gender].append(f"{first} {last}")
    return names
```

Such that:

    >>> group_by_gender(sample_data)
    
    ... {'F': ['Quintina Firle', 'Lena Stockley', 'Margaux Sirr'], 'M': ['Jesse Nunson', 'Taylor Alishoner']}

### Split and Join

A single line with split and join can split a multiline string

    >>> text = """one <br>
    ... two <br>
    ... three"""

into

    >>> 'one|two|three'

Write down an instruction with one or more lines that does this

```python
def join_n_split(text):
    text = '|'.join(text.split('\n'))
    return(text)


# TEST
textex = """one
two
three"""

join_n_split(textex)
```

### String to Tuple

Outline a function which takes a string as input and returns a tuple with two objects
* a modified string with vowels replaced with a star '*'
* the number of replacements

such that

    >>> strip_vowels('hello world')
    ('h*ll* w*rld', 3)

```python
def strip_vowels(text):
    new = []
    count = 0
    for c in text:
        if c.lower() in 'eiyaou':
            new.append('*')
            count += 1
        else:
            new.append(c)
    modified = ''.join(new)
    return (modified, count)
```

### Role of binary operators "//" and "%"

What is the role of the binary operators // and % in Python arithmetics?

    >>> Truncated integer division and modulus (remainder)

### Calling a Python script from the command line

When calling a Python script from the command line

    >>> $ python script.py arg1 arg2

The arguments are saved in a specific data structure of a specific module What is the type and name of the data structure and which module?


__Answer:__

    >>> They are colleced in the list "sys.argv"

### An example of an important usage of the variable `__name__`

Give an example of an important usage of the variable `__name__` in Python scripts

__Answer:__

To filter out code that should not run during import of a module, that only runs when the script is the main program

    >>> if __name__ == "__main__":
    ...

### Name three built-in numerical data types in Python.

`int, float, bool, complex`

### What is the difference between ``=`` and ``==`` in Python

`The first is an assignment operator the second an equality operator`

### Main difference between ``list`` and ``tuple``

The two most common container types for an ordered sequence of objects are list and tuple. What is the main difference between them.

`A list is mutable, a tuple immutable (can not be changed)`

### Writing an ``f-string`` which incorporates the value of a variable v into a string

    >>> v = 42
    >>> text = f"The value of v is {v}." # HIDE
    >>> print(text)
    The value of v is 42.

### A ``dictionary`` is a mapping of key-value pairs. Keys have a restriction...

Keys have a restriction that makes one of the following statements pass:
* key = ()
* key = []
* key = {}

Consider:

    >>> d = {}
    >>> key = () # HIDE
    >>> d[key] = ’somevalue’

Which value of key will pass and why?

`(), because it is immutable`

### Outline a Python function that calculates the ``factorial`` of a non-negative integer `n`

The accumulator pattern is a common design pattern using a for-loop for generating a single value from a given sequence of values. 

Outline a Python function that calculates the factorial of a non-negative integer n, such that:

    >>> f(0)
    1
    >>> f(1)
    1
    >>> f(4)
    24

Solution:

```python
def f(n):
    p = 1
    for i in range(1, n + 1):
        p = p * i
    return p
```

### The `id()` function

The id function takes an object as input and returns a unique number for
that object, its identity, as output. What is the output of the last line,
True of False? This question is about the nature of assignment in Python
so think carefully.

    >>> a = 42
    >>> b = 42
    >>> id(a) == id(b)

__Answer:__

    True

### Output of the `append` operation

    >>> l1 = [1, 2]; l2 = [3, 4]
    >>> l1.append(l2)
    >>> print(l1)

Answer:
`[1, 2, [3, 4]]`

### Outline a ``function`` that takes arbitrary number of input arguments (positional) and returns a list with arguments in ``reverse order``.

Such that:

    >>> f(1)
    [1]
    >>> f(1, 2)
    [2, 1]

Solution:

```python
def f(*args):
    return list(reversed(args))
```

### General info regarding variables. Which are true..?

Literal is a notation specific for a language to represent an object of a given datatype and value.

Which of the following are true:

- 1 represents an object of type int, an integer with the numerical value one.
- 1.0 represents an object of type float, an floating point number with the numerical value one.
- "1.0" represents an object of type str, a string with three characters
- [] represents an object of type list, an empty list
- () represents an object of type tuple, an empty tuple
- {} represents an object of type dict, an empty dictionary

__Answer: all are true__

### Reorder to print correct string

Reorder the lines (write numbers 1-4 next to the lines) to return the output 'data'

    >>> a = 'd'             # 1
    >>> a += 'a'-           # 2
    >>> d = a + 't' + 'a'   # 3
    >>> print(d)            # 4

Correct order (1,3,2,4):

    >>> a = 'd'             # 1
    >>> d = a + 't' + 'a'   # 3
    >>> a += 'a'-           # 2
    >>> print(d)            # 4
